# Trabalho Prático Parte 1

```python
import numpy as np
import math
```

## Questão 1

Calcule as abscissas $t_i$ e os pesos $w_i$ via a matriz de Jacobi.

<!-- grade_id: cell-ad6b84cbfcb7a919, points: 0, grade: false, solution: true -->
```python
def jacobi(n):    
    # Iniciar a matriz J
    # Calcular autovalores e autovetores
    # Ordenar em ordem decrescente do autovalor
    # Calcular w
    # Retornar t e w como vetores de dimensões 1 x n
    #return t, w
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1f4112991a3d4d78, points: 1, grade: true, solution: false -->
```python
ans2 = ([ 0.57735027, -0.57735027], [0.9999999999999998, 0.9999999999999998])
ans3 = ([ 7.74596669e-01, -4.78946310e-17, -7.74596669e-01], [0.5555555555555551, 0.8888888888888894, 0.5555555555555556])
ans4 = ([ 0.86113631,  0.33998104, -0.33998104, -0.86113631], [0.3478548451374544, 0.652145154862546, 0.6521451548625463, 0.3478548451374541])
assert np.allclose(ans2, jacobi(2))
assert np.allclose(ans3, jacobi(3))
assert np.allclose(ans4, jacobi(4))
```

Inicie os vetores ```difTab2```, ```difTab3``` e ```difTab4``` com as diferenças entre os $t$s e $w$s encontrados para $n=2,3,4$ com os valores tabelados.

<!-- grade_id: cell-e95dcb741df9b88b, points: 0, grade: false, solution: true -->
```python
#difTab3[0][0] guarda t[0] - t0 tabelado pra n=3
#difTab3[0][1] guarda t[1] - t1 tabelado pra n=3
#difTab3[0][2] guarda t[2] - t2 tabelado pra n=3
#
#difTab3[1][0] guarda w[0] - w0 tabelado pra n=3
#difTab3[1][1] guarda w[1] - w1 tabelado pra n=3
#difTab3[1][2] guarda w[2] - w2 tabelado pra n=3
#...

#difTab2 = np.zeros((2,2))
#difTab3 = np.zeros((2,3))
#difTab4 = np.zeros((2,4))

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print("difTab2: ", difTab2, "\n")
print("difTab3: ", difTab3, "\n")
print("difTab4: ", difTab4, "\n")
```

<!-- grade_id: cell-40446e53fce845a8, points: 1, grade: true, solution: false -->
```python

```

## Questão 2

```python
print(jacobi(2))
print(jacobi(3))
print(jacobi(4))
```

Porque para $n=2,3,4$ a soma dos pesos é igual a 2?

<!-- grade_id: cell-9df27618cca13eef, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3

Escreva uma função para calcular a integral pela quadratura de Gauss-Legendre.

<!-- grade_id: cell-c9cd894ad4f159d8, points: 0, grade: false, solution: true -->
```python
# y é um vetor 1 x n em que y[i]=f(x[i])
# w é o vetor 1 x n obtido com a função jacobi
def gaussLegendre(y, n, a, b, w):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Calcule $\int_0^2(x^5-x^3+x-1)dx$ usando $n=2$ na função de Gauss-Legendre que você criou.

<!-- grade_id: cell-24d3afa4b8c7f29c, points: 0, grade: false, solution: true -->
```python
#intGL2 = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bfaf2a30118894cc, points: 1, grade: true, solution: false -->
```python
assert np.isclose(5.777777777777776, intGL2)
```

Calcule $\int_0^2(x^5-x^3+x-1)dx$ usando $n=3$ na função de Gauss-Legendre que você criou.

<!-- grade_id: cell-3b1979a4ffb511c4, points: 0, grade: false, solution: true -->
```python
#intGL3 = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-46ac0af9803066a3, points: 1, grade: true, solution: false -->
```python
assert np.isclose(6.6666666666666625, intGL3)
```

Calcule $\int_0^2(x^5-x^3+x-1)dx$ usando $n=4$ na função de Gauss-Legendre que você criou.

<!-- grade_id: cell-30e1ea1ede32837a, points: 0, grade: false, solution: true -->
```python
#intGL4 = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-706b2dcd412eebb8, points: 1, grade: true, solution: false -->
```python
assert np.isclose(6.666666666666674, intGL4)
```

Explique os resultados obtidos para $\int_0^2 (x^5-x^3+x-1) dx$ com $n = 2,3,4$.

<!-- grade_id: cell-56e0ae9d0d3cf65d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 4

Qual o valor exato de $\int_0^\pi \sin(x) dx$?

<!-- grade_id: cell-9f1a2fa17896271f, points: 0, grade: false, solution: true -->
```python
#exact = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e9a7014e37ceb68d, points: 1, grade: true, solution: false -->
```python

```

Inicie o vetor ```difExac``` com as diferenças entre os valores encontrados para $n=3,4,5$ e o valor exato.

<!-- grade_id: cell-249135c994f6e4c3, points: 0, grade: false, solution: true -->
```python
#difExac[0] guarda exact - resultado pra n=3
#difExac[1] guarda exact - resultado pra n=4
#difExac[2] guarda exact - resultado pra n=5
#difExac =

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a183019bae6ee8ed, points: 1, grade: true, solution: false -->
```python

```

```python
print(difExac)
```

## Questão 5

Seja $f(x)=e^x$. Avalie $f(x)$ em $x_0=-\frac{\sqrt{3}}{3}$ e $x_1=\frac{\sqrt{3}}{3}$.

<!-- grade_id: cell-f768f7697710eb50, points: 0, grade: false, solution: true -->
```python
#y[0] guarda f(x0)
#y[1] guarda f(x1)
#y = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-81055d42ea8b7cd9, points: 1, grade: true, solution: false -->
```python

```

```python
print(y)
```

Determine o polinômio linear $p(x)$ que passa pelos pontos ($x_0$,y[0]) e ($x_1$,y[1]).

<!-- grade_id: cell-6458960881d9f29c, points: 0, grade: false, solution: true -->
```python
#p[0] guarda o coeficiente de x^0
#p[1] guarda o coeficiente de x^1
#p = 

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(p)
```

<!-- grade_id: cell-5418addb6d87c943, points: 1, grade: true, solution: false -->
```python

```

Calcule analiticamente $\int_{-1}^1p(x)dx$.

<!-- grade_id: cell-3d8da7afbb17b965, points: 0, grade: false, solution: true -->
```python
#integAn =

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c902cc1d9b925d3e, points: 1, grade: true, solution: false -->
```python

```

Agora calcule $\int_{-1}^1e^xdx$ usando $n=2$ na função de Gauss-Legendre que você criou.

<!-- grade_id: cell-5c111e25fec6a923, points: 0, grade: false, solution: true -->
```python
#integGL =

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7c5458ec4d828f70, points: 1, grade: true, solution: false -->
```python

```

```python
print(integGL - integAn)
```

Explique os resultados.

<!-- grade_id: cell-14eaa15bb1e4fffc, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE